
// I AM NOT DONE

#include <stdio.h>
#include <stdlib.h>

int main(void) {
    printf("Hello !\n");
    printf("This exercise compiles successfully. The remaining exercises contain a compiler\n");
    printf("or logic error. The central concept behind Clings is to fix these errors and\n");
    printf("solve the exercises. Good luck!\n");
    printf("\n");
    printf("To pass to the next kata, simply remove the comment // I AM NOT DONE at the\n");
    printf("top of the kata file.\n");
    printf("The source for this exercise is in `katas/00_intro/intro.c`. Have a look!");

    return EXIT_SUCCESS;
}
